<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Expressions &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="_static/video.css" />
    <link rel="stylesheet" type="text/css" href="_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="_static/reveal.css" />
    <link rel="stylesheet" type="text/css" href="_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="_static/js_lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="_static/accessibility.css" />
    <link rel="stylesheet" type="text/css" href="_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="_static/poll.css" />
    <link rel="stylesheet" href="_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="./" src="_static/documentation_options.js"></script>
    <script type="text/javascript" src="_static/runestonebase.js"></script>
    <script type="text/javascript" src="_static/jquery.js"></script>
    <script type="text/javascript" src="_static/underscore.js"></script>
    <script type="text/javascript" src="_static/doctools.js"></script>
    <script type="text/javascript" src="_static/matrixeq.js"></script>
    <script type="text/javascript" src="_static/clickable.js"></script>
    <script type="text/javascript" src="_static/timedclickable.js"></script>
    <script type="text/javascript" src="_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="_static/fitb.js"></script>
    <script type="text/javascript" src="_static/timedfitb.js"></script>
    <script type="text/javascript" src="_static/runestonevideo.js"></script>
    <script type="text/javascript" src="_static/codemirror.js"></script>
    <script type="text/javascript" src="_static/xml.js"></script>
    <script type="text/javascript" src="_static/css.js"></script>
    <script type="text/javascript" src="_static/htmlmixed.js"></script>
    <script type="text/javascript" src="_static/javascript.js"></script>
    <script type="text/javascript" src="_static/webglinteractive.js"></script>
    <script type="text/javascript" src="_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="_static/Blob.js"></script>
    <script type="text/javascript" src="_static/reveal.js"></script>
    <script type="text/javascript" src="_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="_static/parsons.js"></script>
    <script type="text/javascript" src="_static/timedparsons.js"></script>
    <script type="text/javascript" src="_static/animationbase.js"></script>
    <script type="text/javascript" src="_static/mchoice.js"></script>
    <script type="text/javascript" src="_static/timedmc.js"></script>
    <script type="text/javascript" src="_static/timed.js"></script>
    <script type="text/javascript" src="_static/shortanswer.js"></script>
    <script type="text/javascript" src="_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="_static/pytutor.js"></script>
    <script type="text/javascript" src="_static/codelens.js"></script>
    <script type="text/javascript" src="_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="_static/bookfuncs.js"></script>
    <script type="text/javascript" src="_static/python.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="_static/skulpt.min.js"></script>
    <script type="text/javascript" src="_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="_static/activecode.js"></script>
    <script type="text/javascript" src="_static/clike.js"></script>
    <script type="text/javascript" src="_static/timed_activecode.js"></script>
    <script type="text/javascript" src="_static/datafile.js"></script>
    <script type="text/javascript" src="_static/showEval.js"></script>
    <script type="text/javascript" src="_static/dragndrop.js"></script>
    <script type="text/javascript" src="_static/timeddnd.js"></script>
    <script type="text/javascript" src="_static/poll.js"></script>
    <script type="text/javascript" src="_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="_static/jquery-fix.js"></script>
    <script type="text/javascript" src="_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="_static/waypoints.min.js"></script>
    <script type="text/javascript" src="_static/user-highlights.js"></script>
    <script type="text/javascript" src="_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="genindex.html" />
    <link rel="search" title="Search" href="search.html" />
    <link rel="next" title="String Basics" href="str_basics.html" />
    <link rel="prev" title="Input and Output" href="io.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'python_problems';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'python_problems';
  eBookConfig.acDefaultLanguage = 'python3' ? 'python3' : 'python'
  eBookConfig.runestone_version = '3.2.15';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="index.html">Problems in Python</a>
      </div>
    </div>

     

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/python_problems/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/python_problems/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=python_problems&page=expressions'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Expressions</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Input and Output" data-toggle="tooltip">
    <a href="io.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - String Basics' data-toggle="tooltip" >              
    <a href="str_basics.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>
      
    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="expressions">
<h1>Expressions<a class="headerlink" href="#expressions" title="Permalink to this headline">¶</a></h1>
<div style='display:none;'>
<div data-childcomponent="expr1_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr1_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
x = 4.2
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="expr1_q" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <div id='expr1_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr1" class="runestone explainer ac_section alert alert-warning">
Write the following statement in legal python <code class="docutils literal notranslate"><span class="pre">y</span> <span class="pre">=</span> <span class="pre">x^2</span> <span class="pre">+</span> <span class="pre">2x</span> <span class="pre">+</span> <span class="pre">10</span></code>, assuming that
variable <code class="docutils literal notranslate"><span class="pre">x</span></code> already refers to a value.
<textarea data-component="activecode" id=expr1 data-lang="python" 
     data-include="expr1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
       self.assertAlmostEqual(x, 4.2)
       self.assertAlmostEqual(y, 36.04)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
The exponentiation operator is <code class="docutils literal notranslate"><span class="pre">**</span></code>.  Also, you cannot put <code class="docutils literal notranslate"><span class="pre">2</span></code> and <code class="docutils literal notranslate"><span class="pre">x</span></code> directly next
to each other to do multiplication.  You have to use <code class="docutils literal notranslate"><span class="pre">*</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr1_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr1_a data-lang="python" 
     data-include="expr1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
y = x ** 2 + 2 * x + 10
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="expr2_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr2_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
miles = 42.5
gallons = 1.3
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="expr2_q" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    <div id='expr2_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr2" class="runestone explainer ac_section alert alert-warning">
Given two variables, <code class="docutils literal notranslate"><span class="pre">miles</span></code> and <code class="docutils literal notranslate"><span class="pre">gallons</span></code>, write an assignment statement
that creates the variable <code class="docutils literal notranslate"><span class="pre">mpg</span></code> set to the computed number of miles per gallon.
<textarea data-component="activecode" id=expr2 data-lang="python" 
     data-include="expr2_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
       self.assertAlmostEqual(mpg, miles / gallons)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use the true division operator <code class="docutils literal notranslate"><span class="pre">/</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr2_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr2_a data-lang="python" 
     data-include="expr2_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
mpg = miles / gallons
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="expr3_q" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='expr3_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr3" class="runestone explainer ac_section alert alert-warning">
Create a variable <code class="docutils literal notranslate"><span class="pre">remainder</span></code> that refers to the remainder when dividing 17 by 5.
<textarea data-component="activecode" id=expr3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
       self.assertAlmostEqual(remainder, 17 % 5)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
To get the remainder when dividing two numbers, you need to use the <em>modulo</em> operator <code class="docutils literal notranslate"><span class="pre">%</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr3_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr3_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
remainder = 17 % 5
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="expr4_q" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    <div id='expr4_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr4" class="runestone explainer ac_section alert alert-warning">
<p>Write a small program that asks the user for two numbers and prints out the remainder
when dividing the first by the second.  An example run should look like this:</p>
<p>Enter the numerator: <code class="docutils literal notranslate"><span class="pre">22</span></code></p>
<p>Enter the denominator: <code class="docutils literal notranslate"><span class="pre">4</span></code></p>
<p>The remainder when dividing is 2</p>

<textarea data-component="activecode" id=expr4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
To get the remainder when dividing two numbers, you need to use the <em>modulo</em> operator <code class="docutils literal notranslate"><span class="pre">%</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr4_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr4_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
numerator = int(input('Enter the numerator: '))
denominator = int(input('Enter the denominator: '))
print('The remainder when dividing is', numerator % denominator)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="expr5_q" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='expr5_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr5" class="runestone explainer ac_section alert alert-warning">
You are a bean counter who counts beans and puts them into piles of 16 beans each.
Your manager gives you some number of beans and asks you to tell him how many full piles
that will produce.  Write a program that asks for the number of beans and prints
out the number of full piles.
<textarea data-component="activecode" id=expr5 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use integer division, which divides two numbers, and then “throws out” the fractional part.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr5_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr5_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
beans = int(input('Enter # of beans: '))
print(beans // 16)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="expr6_q" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    <div id='expr6_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr6" class="runestone explainer ac_section alert alert-warning">
<p>Remainder and integer division can often be used together.
Alter the program from the last exercise to print out not only the number of
full piles of beans, but also the number of beans that remain.
E.g., a run of your program might look like this:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="c1"># of beans: 99</span>
<span class="mi">99</span> <span class="n">beans</span> <span class="ow">is</span> <span class="mi">6</span> <span class="n">piles</span> <span class="k">with</span> <span class="mi">3</span> <span class="n">remaining</span>
</pre></div>
</div>

<textarea data-component="activecode" id=expr6 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Use integer division, which divides two numbers, and then “throws out” the fractional part,
and modulo division, which divides two numbers, keeping the remainder.
You might find it useful to create 3 variables: <code class="docutils literal notranslate"><span class="pre">beans</span></code> is the count of beans the user entered,
<code class="docutils literal notranslate"><span class="pre">piles</span></code> is the computed number of piles, and <code class="docutils literal notranslate"><span class="pre">remaining</span></code> is the computed remainder.</p>
<p>Also, remember that you cannot add together integers and strings. If you are going to
concatenate, both items must be strings.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr6_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr6_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
beans = int(input('Enter # of beans: '))
piles = beans // 16
remaining = beans % 16
print(beans, 'beans is', piles, 'piles with', remaining, 'remaining')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="expr7_q" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='expr7_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr7" class="runestone explainer ac_section alert alert-warning">
<p>The <em>discriminant</em> of the quadratic formula is the expression:
<code class="docutils literal notranslate"><span class="pre">b^2</span> <span class="pre">-</span> <span class="pre">4ac</span></code></p>
<p>Write a program to ask the user for a value for <code class="docutils literal notranslate"><span class="pre">a</span></code>, a value for <code class="docutils literal notranslate"><span class="pre">b</span></code>,
and a value for <code class="docutils literal notranslate"><span class="pre">c</span></code>, and then prints the discriminant.</p>

<textarea data-component="activecode" id=expr7 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Don’t forget that you cannot just put variables <code class="docutils literal notranslate"><span class="pre">a</span></code> and <code class="docutils literal notranslate"><span class="pre">c</span></code> next to each
other.  You have to use <code class="docutils literal notranslate"><span class="pre">*</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr7_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr7_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
a = int(input('Enter a: '))
b = int(input('Enter b: '))
c = int(input('Enter c: '))
print(b * b - 4 * a * c)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="expr8_q" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    <div id='expr8_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr8" class="runestone explainer ac_section alert alert-warning">
<p>Write a short program that asks the user to enter an integer, and then prints
out the last digit of the number that was entered.  E.g.,:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Enter</span> <span class="n">an</span> <span class="n">integer</span><span class="p">:</span> <span class="mi">162</span>
<span class="n">The</span> <span class="n">last</span> <span class="n">digit</span> <span class="ow">is</span> <span class="mi">2</span>
</pre></div>
</div>

<textarea data-component="activecode" id=expr8 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
The last digit of a number is the remainder when dividing by 10.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr8_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr8_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
anInt = int(input('Enter an integer: '))
lastDigit = anInt % 10
print('The last digit is', lastDigit)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="expr9_q" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='expr9_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="expr9" class="runestone explainer ac_section alert alert-warning">
<p>Write a short program that computes a total restaurant bill, including a 15%
tip.  E.g.,:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>Enter the amount on the bill: 30.00
The 15% tip amount is: $4.5
The total is: $34.5
</pre></div>
</div>

<textarea data-component="activecode" id=expr9 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
To read in a decimal value (a <code class="docutils literal notranslate"><span class="pre">floating</span> <span class="pre">point</span> <span class="pre">number</span></code>), use <code class="docutils literal notranslate"><span class="pre">variable</span> <span class="pre">=</span> <span class="pre">float(input('prompt'))</span></code>.
Use multiple variables and create your program one line at a time.  Remember that
15% is 0.15. Don’t worry about always having 2 digits after the decimal point.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="expr9_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=expr9_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
subtotal = float(input('Enter the amount on the bill: '))
tip = subtotal * 0.15
print("The 15% tip amount is: $", tip)
total = subtotal + tip
print('The total is: $', total)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    </div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Input and Output' data-toggle="tooltip">
    <a href="io.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="io.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - String Basics' data-toggle="tooltip" >
    <a href="str_basics.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="str_basics.html" >Next Section - String Basics</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2019 vtn2.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 3.2.15.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>