<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>For Loops &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="_static/video.css" />
    <link rel="stylesheet" type="text/css" href="_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="_static/reveal.css" />
    <link rel="stylesheet" type="text/css" href="_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="_static/js_lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="_static/accessibility.css" />
    <link rel="stylesheet" type="text/css" href="_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="_static/poll.css" />
    <link rel="stylesheet" href="_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="./" src="_static/documentation_options.js"></script>
    <script type="text/javascript" src="_static/runestonebase.js"></script>
    <script type="text/javascript" src="_static/jquery.js"></script>
    <script type="text/javascript" src="_static/underscore.js"></script>
    <script type="text/javascript" src="_static/doctools.js"></script>
    <script type="text/javascript" src="_static/matrixeq.js"></script>
    <script type="text/javascript" src="_static/clickable.js"></script>
    <script type="text/javascript" src="_static/timedclickable.js"></script>
    <script type="text/javascript" src="_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="_static/fitb.js"></script>
    <script type="text/javascript" src="_static/timedfitb.js"></script>
    <script type="text/javascript" src="_static/runestonevideo.js"></script>
    <script type="text/javascript" src="_static/codemirror.js"></script>
    <script type="text/javascript" src="_static/xml.js"></script>
    <script type="text/javascript" src="_static/css.js"></script>
    <script type="text/javascript" src="_static/htmlmixed.js"></script>
    <script type="text/javascript" src="_static/javascript.js"></script>
    <script type="text/javascript" src="_static/webglinteractive.js"></script>
    <script type="text/javascript" src="_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="_static/Blob.js"></script>
    <script type="text/javascript" src="_static/reveal.js"></script>
    <script type="text/javascript" src="_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="_static/parsons.js"></script>
    <script type="text/javascript" src="_static/timedparsons.js"></script>
    <script type="text/javascript" src="_static/animationbase.js"></script>
    <script type="text/javascript" src="_static/mchoice.js"></script>
    <script type="text/javascript" src="_static/timedmc.js"></script>
    <script type="text/javascript" src="_static/timed.js"></script>
    <script type="text/javascript" src="_static/shortanswer.js"></script>
    <script type="text/javascript" src="_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="_static/pytutor.js"></script>
    <script type="text/javascript" src="_static/codelens.js"></script>
    <script type="text/javascript" src="_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="_static/bookfuncs.js"></script>
    <script type="text/javascript" src="_static/python.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="_static/skulpt.min.js"></script>
    <script type="text/javascript" src="_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="_static/activecode.js"></script>
    <script type="text/javascript" src="_static/clike.js"></script>
    <script type="text/javascript" src="_static/timed_activecode.js"></script>
    <script type="text/javascript" src="_static/datafile.js"></script>
    <script type="text/javascript" src="_static/showEval.js"></script>
    <script type="text/javascript" src="_static/dragndrop.js"></script>
    <script type="text/javascript" src="_static/timeddnd.js"></script>
    <script type="text/javascript" src="_static/poll.js"></script>
    <script type="text/javascript" src="_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="_static/jquery-fix.js"></script>
    <script type="text/javascript" src="_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="_static/waypoints.min.js"></script>
    <script type="text/javascript" src="_static/user-highlights.js"></script>
    <script type="text/javascript" src="_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="genindex.html" />
    <link rel="search" title="Search" href="search.html" />
    <link rel="prev" title="Range Function" href="range.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'python_problems';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'python_problems';
  eBookConfig.acDefaultLanguage = 'python3' ? 'python3' : 'python'
  eBookConfig.runestone_version = '3.2.15';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="index.html">Problems in Python</a>
      </div>
    </div>

     

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/python_problems/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/python_problems/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=python_problems&page=for_loops'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">For Loops</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Range Function" data-toggle="tooltip">
    <a href="range.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
    <li class="divider-vertical"></li>
  

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>
      
    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="for-loops">
<h1>For Loops<a class="headerlink" href="#for-loops" title="Permalink to this headline">¶</a></h1>
<div style='display:none;'>
<div data-childcomponent="for1_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for1_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
groceries = ['milk', 'bread', 'eggs', 'chocolate (dark)']
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for1_q" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <blockquote>
<div><div id='for1_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for1" class="runestone explainer ac_section alert alert-warning">
<p>Assume a variable <code class="docutils literal notranslate"><span class="pre">groceries</span></code> has been defined to be a
list of grocery items.</p>
<p>Write code to print out each item on its own line.</p>

<textarea data-component="activecode" id=for1 data-lang="python" 
     data-include="for1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>The for loop pattern is:</p>
<p><code class="docutils literal notranslate"><span class="pre">for</span> <span class="pre">item</span> <span class="pre">in</span> <span class="pre">list:</span></code></p>
<blockquote>
<div><code class="docutils literal notranslate"><span class="pre">loop-body</span></code></div></blockquote>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for1_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for1_a data-lang="python" 
     data-include="for1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for item in groceries:
    print(item)
</textarea>
</div>
</div>
    </div>

</div></blockquote>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for2_q" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    <div id='for2_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for2" class="runestone explainer ac_section alert alert-warning">
<p>Assume a variable <code class="docutils literal notranslate"><span class="pre">groceries</span></code> has been defined to be a
list of grocery items.</p>
<p>Write code to print out all items in the list on separate
lines, prefixed by the number of the item in the list.
E.g.,</p>
<ol class="arabic simple">
<li>milk</li>
<li>eggs</li>
<li>chips</li>
</ol>

<textarea data-component="activecode" id=for2 data-lang="python" 
     data-include="for1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
You will have to use an index-based loop, with
<code class="docutils literal notranslate"><span class="pre">range(len(groceries))</span></code> and indexing groceries as <code class="docutils literal notranslate"><span class="pre">groceries[idx]</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for2_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for2_a data-lang="python" 
     data-include="for1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in range(len(groceries)):
    print(str(idx + 1) + '. ' + groceries[idx])
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for3_q" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='for3_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for3" class="runestone explainer ac_section alert alert-warning">
<p>Assume a variable <code class="docutils literal notranslate"><span class="pre">groceries</span></code> has been defined to be a
list of grocery items.</p>
<p>Write code to print out all items in the list on a single
line, separated by commas.  It is OK to have a comma after
the last item.</p>

<textarea data-component="activecode" id=for3 data-lang="python" 
     data-include="for1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Remember that you can use <code class="docutils literal notranslate"><span class="pre">print(...,</span> <span class="pre">end='')</span></code> to keep the cursor
on the same line.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for3_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for3_a data-lang="python" 
     data-include="for1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for item in groceries:
    print(item, end=', ')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="for4_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for4_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
n = 10
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for4_q" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    <div id='for4_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for4" class="runestone explainer ac_section alert alert-warning">
Write a for loop to sum up all numbers from 0 to <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">-</span> <span class="pre">1</span></code> into
a variable <code class="docutils literal notranslate"><span class="pre">sum</span></code>.  Assume <code class="docutils literal notranslate"><span class="pre">n</span></code> is already defined.
<textarea data-component="activecode" id=for4 data-lang="python" 
     data-include="for4_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(sum, 45)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Your code should have this format:</p>
<ul class="simple">
<li>initialize sum to 0.</li>
<li><dl class="first docutils">
<dt>for loop, using <code class="docutils literal notranslate"><span class="pre">range</span></code>.</dt>
<dd><ul class="first last">
<li>add the value to sum</li>
</ul>
</dd>
</dl>
</li>
</ul>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for4_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for4_a data-lang="python" 
     data-include="for4_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
sum = 0
for num in range(0, n):
    sum = sum + num
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="for5_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for5_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
name = 'Abraham Lincoln'
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for5_q" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='for5_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for5" class="runestone explainer ac_section alert alert-warning">
Write a for loop to print out each letter in a str <code class="docutils literal notranslate"><span class="pre">name</span></code>
on one line with a space between.
<textarea data-component="activecode" id=for5 data-lang="python" 
     data-include="for5_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>A string is iterable, so you can use it right in the for loop:</p>
<p>e.g.,</p>
<p><code class="docutils literal notranslate"><span class="pre">for</span> <span class="pre">char</span> <span class="pre">in</span> <span class="pre">str:</span></code></p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for5_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for5_a data-lang="python" 
     data-include="for5_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for char in name:
    print(char, end=' ')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="for6_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for6_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
name = 'Abraham Lincoln'
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for6_q" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    <div id='for6_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for6" class="runestone explainer ac_section alert alert-warning">
Write a for loop to iterate over a string <code class="docutils literal notranslate"><span class="pre">name</span></code> and print
every other letter  – i.e., the letters with indices
which are even.  That is, only the 0th, 2nd, 4th, etc., letters.
<textarea data-component="activecode" id=for6 data-lang="python" 
     data-include="for6_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Use an index-based loop, since you will need to know if you are
looking at an even or odd index.</p>
<p>Use an if statement in the for loop to determine if the index is even.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for6_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for6_a data-lang="python" 
     data-include="for6_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in range(len(name)):
    if idx % 2 == 0:
        print(name[idx])
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for7_q" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='for7_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for7" class="runestone explainer ac_section alert alert-warning">
Write code to create a list <code class="docutils literal notranslate"><span class="pre">rands</span></code> that contains
20 random numbers from 0 to 100 inclusive.
<textarea data-component="activecode" id=for7 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
from random import randint
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(len(rands), 20)
       for item in rands:
           self.assertTrue(item >= 0 and item <= 100)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Above the for loop, create the variable <code class="docutils literal notranslate"><span class="pre">rands</span></code> to be an empty list.</p>
<p>Call <code class="docutils literal notranslate"><span class="pre">randint(0,</span> <span class="pre">100)</span></code> to get random integers from 0 to 100,
inclusive.</p>
<p>Use <code class="docutils literal notranslate"><span class="pre">append()</span></code> to add a random number to the list.</p>
<p>Use <code class="docutils literal notranslate"><span class="pre">range(20)</span></code> to iterate 20 times.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for7_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for7_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
from random import randint
rands = []
for idx in range(20):
    rands.append(randint(0, 100))
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for8_q" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    <div id='for8_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for8" class="runestone explainer ac_section alert alert-warning">
Write a for loop to create a list <code class="docutils literal notranslate"><span class="pre">tenths</span></code> containing
numbers 0.0, 0.1, 0.2, 0.3, 0.4, …, 9.9
<textarea data-component="activecode" id=for8 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(len(tenths), 100)
       self.assertAlmostEqual(tenths[0], 0.0)
       self.assertAlmostEqual(tenths[99], 9.9)
       self.assertAlmostEqual(tenths[50], 5.0)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Above the for loop, create the variable <code class="docutils literal notranslate"><span class="pre">tenths</span></code> to be an empty list.</p>
<p>There will be 100 numbers in the list, so use <code class="docutils literal notranslate"><span class="pre">range(100)</span></code> to iterate 100 times.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for8_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for8_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
tenths = []
for i in range(100):
    tenths.append(i / 10)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="for9_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for9_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
numbers = [ 44, 3, 12, 22, 71, 47.3, 99.22, 105.1, 3, 7]
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for9_q" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='for9_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for9" class="runestone explainer ac_section alert alert-warning">
Given a list <code class="docutils literal notranslate"><span class="pre">numbers</span></code>, write code to print <code class="docutils literal notranslate"><span class="pre">Found</span> <span class="pre">it!</span></code> if
the list contains the number 7. Do not use the <code class="docutils literal notranslate"><span class="pre">in</span></code> operator.  Use a for loop,
with an if statement in it.
<textarea data-component="activecode" id=for9 data-lang="python" 
     data-include="for9_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use item-based iteration.  Put a simple if statement in the for loop.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for9_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for9_a data-lang="python" 
     data-include="for9_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for num in numbers:
    if num == 7:
        print("Found it!")
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="for10_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for10_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
numbers = [ 44, 3, 12, 22, 71, 47.3, 99.22, 105.1, 3, 7, 50]
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for10_q" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    <div id='for10_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for10" class="runestone explainer ac_section alert alert-warning">
Given a list <code class="docutils literal notranslate"><span class="pre">numbers</span></code>, write code to create two lists, <code class="docutils literal notranslate"><span class="pre">bignums</span></code> and <code class="docutils literal notranslate"><span class="pre">littlenums</span></code>,
by adding each number in <code class="docutils literal notranslate"><span class="pre">numbers</span></code> to <code class="docutils literal notranslate"><span class="pre">bignums</span></code> if it is greater than 50, and adding
all the other numbers to <code class="docutils literal notranslate"><span class="pre">littlenums</span></code>.
<textarea data-component="activecode" id=for10 data-lang="python" 
     data-include="for10_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(len(littlenums) + len(bignums), len(numbers))
       self.assertTrue(50 in littlenums, "50 should be in littlenums")
       self.assertEqual(len(littlenums), 8)
       self.assertTrue(99.22 in bignums)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Create the resulting lists before the for loop.</p>
<p>Use item-based iteration with an if-else in it.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for10_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for10_a data-lang="python" 
     data-include="for10_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
bignums = []
littlenums = []
for num in numbers:
    if num > 50:
        bignums.append(num)
    else:
        littlenums.append(num)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="for11_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for11_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
class Block:
    def __init__(self, x, y):
        pass
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for11_q" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    <div id='for11_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for11" class="runestone explainer ac_section alert alert-warning">
<p>This code was recently seen in a project:</p>
<p><code class="docutils literal notranslate"><span class="pre">block_list</span> <span class="pre">=</span> <span class="pre">[Block(20,20),</span> <span class="pre">Block(20,60),</span> <span class="pre">Block(20,100),</span> <span class="pre">Block(20,140),</span> <span class="pre">Block(20,180),</span> <span class="pre">Block(60,20),</span> <span class="pre">Block(60,60),</span> <span class="pre">Block(60,100),</span> <span class="pre">Block(60,140),</span> <span class="pre">Block(60,180),</span> <span class="pre">Block(100,20),</span> <span class="pre">Block(100,60),</span> <span class="pre">Block(100,100),</span> <span class="pre">Block(100,140),</span> <span class="pre">Block(100,180),</span> <span class="pre">Block(140,20),</span> <span class="pre">Block(140,60),</span> <span class="pre">Block(140,100),</span> <span class="pre">Block(140,140),</span> <span class="pre">Block(140,180),</span> <span class="pre">Block(180,20),</span> <span class="pre">Block(180,60),</span> <span class="pre">Block(180,100),</span> <span class="pre">Block(180,140),</span> <span class="pre">Block(180,180),</span> <span class="pre">Block(220,20),</span> <span class="pre">Block(220,60),</span> <span class="pre">Block(220,100),</span> <span class="pre">Block(220,140),</span> <span class="pre">Block(220,180),Block(260,20),</span> <span class="pre">Block(260,60),</span> <span class="pre">Block(260,100),</span> <span class="pre">Block(260,140),</span> <span class="pre">Block(260,180),</span> <span class="pre">Block(300,20),</span> <span class="pre">Block(300,60),</span> <span class="pre">Block(300,100),</span> <span class="pre">Block(300,140),</span> <span class="pre">Block(300,180),</span> <span class="pre">Block(340,20),</span> <span class="pre">Block(340,60),</span> <span class="pre">Block(340,100),</span> <span class="pre">Block(340,140),</span> <span class="pre">Block(340,180),</span> <span class="pre">Block(380,20),</span> <span class="pre">Block(380,60),</span> <span class="pre">Block(380,100),</span> <span class="pre">Block(380,140),</span> <span class="pre">Block(380,180),</span> <span class="pre">Block(420,20),</span> <span class="pre">Block(420,60),</span> <span class="pre">Block(420,100),</span> <span class="pre">Block(420,140),</span> <span class="pre">Block(420,180),</span> <span class="pre">Block(460,20),</span> <span class="pre">Block(460,60),</span> <span class="pre">Block(460,100),</span> <span class="pre">Block(460,140),</span> <span class="pre">Block(460,180),</span> <span class="pre">Block(500,20),</span> <span class="pre">Block(500,60),</span> <span class="pre">Block(500,100),</span> <span class="pre">Block(500,140),</span> <span class="pre">Block(500,180),</span> <span class="pre">Block(540,20),</span> <span class="pre">Block(540,60),</span> <span class="pre">Block(540,100),</span> <span class="pre">Block(540,140),</span> <span class="pre">Block(540,180),</span> <span class="pre">Block(580,20),</span> <span class="pre">Block(580,60),</span> <span class="pre">Block(580,100),</span> <span class="pre">Block(580,140),</span> <span class="pre">Block(580,180),</span> <span class="pre">Block(620,20),</span> <span class="pre">Block(620,60),</span> <span class="pre">Block(620,100),</span> <span class="pre">Block(620,140),</span> <span class="pre">Block(620,180),</span> <span class="pre">Block(660,20),</span> <span class="pre">Block(660,60),</span> <span class="pre">Block(660,100),</span> <span class="pre">Block(660,140),</span> <span class="pre">Block(660,180),</span> <span class="pre">Block(700,20),</span> <span class="pre">Block(700,60),</span> <span class="pre">Block(700,100),</span> <span class="pre">Block(700,140),</span> <span class="pre">Block(700,180)]</span></code></p>
<p>Rewrite this code to use a doubly-nest for loop.</p>

<textarea data-component="activecode" id=for11 data-lang="python" 
     data-include="for11_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
block_list = []
# add your code below
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(len(block_list), 90)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Create the resulting list before the for loops.</p>
<p>Create a for loop to iterate from 20 to 700, by 40s.</p>
<p>Inside the first for loop, write a for loop to iterate from 20 to 180 by 40s.</p>
<p>Inside that for loop, call the <code class="docutils literal notranslate"><span class="pre">Block()</span></code> call, and append the result to your <code class="docutils literal notranslate"><span class="pre">block_list</span></code>.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for11_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for11_a data-lang="python" 
     data-include="for11_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
block_list = []
for vert in range(20, 701, 40):
    for horiz in range(20, 181, 40):
        block_list.append(Block(vert, horiz))
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="for12_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for12_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
groceries = ['milk', 'bread', 'eggs', 'chocolate (dark)', 'oatmeal']
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="for12_q" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    <div id='for12_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for12" class="runestone explainer ac_section alert alert-warning">
<p>Integers are used as indices into <em>sequences</em>, like lists and
strings. So, if you want to access each element in a
sequence at a certain index, you need to use a <code class="docutils literal notranslate"><span class="pre">for</span></code> loop
to iterate over each <em>index</em> appropriate for accessing the
sequence at that index.</p>
<p>To get the indices appropriate for indexing a sequence, use
<code class="docutils literal notranslate"><span class="pre">range(len(sequence))</span></code>.</p>
<p>Replace <code class="docutils literal notranslate"><span class="pre">&lt;1&gt;</span></code> in the following code with a call to
<code class="docutils literal notranslate"><span class="pre">range(len(sequence))</span></code> so that the for loop prints out all indices
useful for indexing into the sequence <code class="docutils literal notranslate"><span class="pre">groceries</span></code>.</p>

<textarea data-component="activecode" id=for12 data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in <1>:
    print(idx)
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<code class="docutils literal notranslate"><span class="pre">len(groceries)</span></code> returns the number of items in the sequence.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for12_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for12_a data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in range(len(groceries)):
    # print out the index for this item in the groceries list.
    print(idx)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for13_q" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    <div id='for13_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for13" class="runestone explainer ac_section alert alert-warning">
<p>Replace <code class="docutils literal notranslate"><span class="pre">&lt;1&gt;</span></code> in the following code with a call to
<code class="docutils literal notranslate"><span class="pre">range(len(sequence))</span></code> so that the for loop uses indexes
useful for indexing into the sequence <code class="docutils literal notranslate"><span class="pre">groceries</span></code>.</p>
<p>Replace <code class="docutils literal notranslate"><span class="pre">&lt;2&gt;</span></code> with code to access the value in sequence
<code class="docutils literal notranslate"><span class="pre">groceries</span></code> at index <code class="docutils literal notranslate"><span class="pre">idx</span></code>.</p>
<p>Your resulting code should print each item in the <code class="docutils literal notranslate"><span class="pre">groceries</span></code>
list, regardless of how long the <code class="docutils literal notranslate"><span class="pre">groceries</span></code> list is.</p>

<textarea data-component="activecode" id=for13 data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in <1>:
    print(<2>)
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<code class="docutils literal notranslate"><span class="pre">len(groceries)</span></code> returns the number of items in the sequence.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for13_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for13_a data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in range(len(groceries)):
    print(groceries[idx])
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for14_q" >
    <ol start=14 class=arabic><li class="alert alert-warning">

    <div id='for14_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for14" class="runestone explainer ac_section alert alert-warning">
Replace <code class="docutils literal notranslate"><span class="pre">&lt;1&gt;</span></code> in the code so that the for loop prints each
item <em>except the last</em> in the list <code class="docutils literal notranslate"><span class="pre">groceries</span></code>.
<textarea data-component="activecode" id=for14 data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in <1>:
    print(groceries[idx])
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<code class="docutils literal notranslate"><span class="pre">len(groceries)</span></code> returns the number of items in the sequence,
but we want to access all items except the last one, so subtract
one from that value.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for14_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for14_a data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in range(len(groceries) - 1):
    print(groceries[idx])
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for15_q" >
    <ol start=15 class=arabic><li class="alert alert-warning">

    <div id='for15_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for15" class="runestone explainer ac_section alert alert-warning">
Replace <code class="docutils literal notranslate"><span class="pre">&lt;1&gt;</span></code> in the code so that the for loop prints each
item <em>except the first and the last</em> in the list <code class="docutils literal notranslate"><span class="pre">groceries</span></code>.
<textarea data-component="activecode" id=for15 data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in <1>:
    print(groceries[idx])
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>To skip the first item, start with index <code class="docutils literal notranslate"><span class="pre">1</span></code>, not the default
of <code class="docutils literal notranslate"><span class="pre">0</span></code>.</p>
<p><code class="docutils literal notranslate"><span class="pre">len(groceries)</span></code> returns the number of items in the sequence,
but we want to access all items except the last one, so subtract
one from that value.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for15_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for15_a data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in range(1, len(groceries) - 1):
    print(groceries[idx])
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="for16_q" >
    <ol start=16 class=arabic><li class="alert alert-warning">

    <div id='for16_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="for16" class="runestone explainer ac_section alert alert-warning">
Replace <code class="docutils literal notranslate"><span class="pre">&lt;1&gt;</span></code> in the code so that the for loop prints each
item in the list <code class="docutils literal notranslate"><span class="pre">groceries</span></code> in reverse order.  Use <code class="docutils literal notranslate"><span class="pre">range()</span></code>
with a step of -1.
<textarea data-component="activecode" id=for16 data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in <1>:
    print(groceries[idx])
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
You need to use all three parameters to <code class="docutils literal notranslate"><span class="pre">range</span></code>.  The <code class="docutils literal notranslate"><span class="pre">start</span></code> value
needs to be the index of the last element of <code class="docutils literal notranslate"><span class="pre">groceries</span></code> – so get
the length of <code class="docutils literal notranslate"><span class="pre">groceries</span></code> and subtract 1.  The <code class="docutils literal notranslate"><span class="pre">stop</span></code> value needs to be
less than 0, so that 0 is the last index that <code class="docutils literal notranslate"><span class="pre">range</span></code> generates.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="for16_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=for16_a data-lang="python" 
     data-include="for12_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
for idx in range(len(groceries) - 1, -1, -1):
    print(groceries[idx])
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    </div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Range Function' data-toggle="tooltip">
    <a href="range.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="range.html" ></a>
  

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2019 vtn2.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 3.2.15.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>