<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>If Statements &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="_static/video.css" />
    <link rel="stylesheet" type="text/css" href="_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="_static/reveal.css" />
    <link rel="stylesheet" type="text/css" href="_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="_static/js_lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="_static/accessibility.css" />
    <link rel="stylesheet" type="text/css" href="_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="_static/poll.css" />
    <link rel="stylesheet" href="_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="./" src="_static/documentation_options.js"></script>
    <script type="text/javascript" src="_static/runestonebase.js"></script>
    <script type="text/javascript" src="_static/jquery.js"></script>
    <script type="text/javascript" src="_static/underscore.js"></script>
    <script type="text/javascript" src="_static/doctools.js"></script>
    <script type="text/javascript" src="_static/matrixeq.js"></script>
    <script type="text/javascript" src="_static/clickable.js"></script>
    <script type="text/javascript" src="_static/timedclickable.js"></script>
    <script type="text/javascript" src="_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="_static/fitb.js"></script>
    <script type="text/javascript" src="_static/timedfitb.js"></script>
    <script type="text/javascript" src="_static/runestonevideo.js"></script>
    <script type="text/javascript" src="_static/codemirror.js"></script>
    <script type="text/javascript" src="_static/xml.js"></script>
    <script type="text/javascript" src="_static/css.js"></script>
    <script type="text/javascript" src="_static/htmlmixed.js"></script>
    <script type="text/javascript" src="_static/javascript.js"></script>
    <script type="text/javascript" src="_static/webglinteractive.js"></script>
    <script type="text/javascript" src="_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="_static/Blob.js"></script>
    <script type="text/javascript" src="_static/reveal.js"></script>
    <script type="text/javascript" src="_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="_static/parsons.js"></script>
    <script type="text/javascript" src="_static/timedparsons.js"></script>
    <script type="text/javascript" src="_static/animationbase.js"></script>
    <script type="text/javascript" src="_static/mchoice.js"></script>
    <script type="text/javascript" src="_static/timedmc.js"></script>
    <script type="text/javascript" src="_static/timed.js"></script>
    <script type="text/javascript" src="_static/shortanswer.js"></script>
    <script type="text/javascript" src="_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="_static/pytutor.js"></script>
    <script type="text/javascript" src="_static/codelens.js"></script>
    <script type="text/javascript" src="_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="_static/bookfuncs.js"></script>
    <script type="text/javascript" src="_static/python.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="_static/skulpt.min.js"></script>
    <script type="text/javascript" src="_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="_static/activecode.js"></script>
    <script type="text/javascript" src="_static/clike.js"></script>
    <script type="text/javascript" src="_static/timed_activecode.js"></script>
    <script type="text/javascript" src="_static/datafile.js"></script>
    <script type="text/javascript" src="_static/showEval.js"></script>
    <script type="text/javascript" src="_static/dragndrop.js"></script>
    <script type="text/javascript" src="_static/timeddnd.js"></script>
    <script type="text/javascript" src="_static/poll.js"></script>
    <script type="text/javascript" src="_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="_static/jquery-fix.js"></script>
    <script type="text/javascript" src="_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="_static/waypoints.min.js"></script>
    <script type="text/javascript" src="_static/user-highlights.js"></script>
    <script type="text/javascript" src="_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="genindex.html" />
    <link rel="search" title="Search" href="search.html" />
    <link rel="prev" title="String Formating" href="string_formatting.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'python_problems';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'python_problems';
  eBookConfig.acDefaultLanguage = 'python3' ? 'python3' : 'python'
  eBookConfig.runestone_version = '3.2.15';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="index.html">Problems in Python</a>
      </div>
    </div>

     

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/python_problems/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/python_problems/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=python_problems&page=if_statements'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">If Statements</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - String Formating" data-toggle="tooltip">
    <a href="string_formatting.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
    <li class="divider-vertical"></li>
  

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>
      
    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="if-statements">
<h1>If Statements<a class="headerlink" href="#if-statements" title="Permalink to this headline">¶</a></h1>
<p><strong>if statements</strong> have this general form:</p>
<p><strong>if</strong> <code class="docutils literal notranslate"><span class="pre">conditional</span></code> :</p>
<blockquote>
<div><strong>body</strong></div></blockquote>
<p>and</p>
<p><strong>if-else statements</strong> have this general form:</p>
<p><strong>if</strong> <code class="docutils literal notranslate"><span class="pre">conditional</span></code> :</p>
<blockquote>
<div><strong>if-body</strong></div></blockquote>
<p><strong>else:</strong></p>
<blockquote>
<div><strong>else-body</strong></div></blockquote>
<p>and</p>
<p><strong>if-elif-else statements</strong> have this general form:</p>
<p><strong>if</strong> <code class="docutils literal notranslate"><span class="pre">conditional</span></code> :</p>
<blockquote>
<div><strong>if-body</strong></div></blockquote>
<p><strong>elif</strong> <code class="docutils literal notranslate"><span class="pre">conditional</span></code> :</p>
<blockquote>
<div><strong>elif-body</strong></div></blockquote>
<p><strong>else:</strong></p>
<blockquote>
<div><strong>else-body</strong></div></blockquote>
<div style='display:none;'>
<div data-childcomponent="ifstmt1_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt1_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
v1 = 4
v2 = 5
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt1_q" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <blockquote>
<div><div id='ifstmt1_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt1" class="runestone explainer ac_section alert alert-warning">
Write an if statement that executes the <strong>body</strong> only if <code class="docutils literal notranslate"><span class="pre">v1</span></code> is less than <code class="docutils literal notranslate"><span class="pre">v2</span></code>.
The <strong>body</strong> should simply contain <code class="docutils literal notranslate"><span class="pre">print('yes')</span></code>.
<textarea data-component="activecode" id=ifstmt1 data-lang="python" 
     data-include="ifstmt1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your if statement
    print('yes')
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Look at the top of this file for a reminder about the formatting of an if statement.</p>
<p>You do not need an else part.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt1_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt1_a data-lang="python" 
     data-include="ifstmt1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if v1 < v2:
    print('yes')
</textarea>
</div>
</div>
    </div>

</div></blockquote>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="ifstmt2_q" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    <div id='ifstmt2_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt2" class="runestone explainer ac_section alert alert-warning">
Write an if statement that prints <code class="docutils literal notranslate"><span class="pre">yes</span></code> if <code class="docutils literal notranslate"><span class="pre">v1</span></code> is greater than <code class="docutils literal notranslate"><span class="pre">v2</span></code> and prints <code class="docutils literal notranslate"><span class="pre">no</span></code> otherwise.
<textarea data-component="activecode" id=ifstmt2 data-lang="python" 
     data-include="ifstmt1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Look at the top of this file for a reminder about the formatting of an if-else statement.</p>
<p>You need an else part.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt2_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt2_a data-lang="python" 
     data-include="ifstmt1_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if v1 > v2:
    print('yes')
else:
    print('no')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt3_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt3_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
v1 = 7
v2 = 5
v3 = 4
v4 = 8
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt3_q" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='ifstmt3_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt3" class="runestone explainer ac_section alert alert-warning">
Write an if statement that prints <code class="docutils literal notranslate"><span class="pre">yes</span></code> if <code class="docutils literal notranslate"><span class="pre">v1</span></code> is less than <code class="docutils literal notranslate"><span class="pre">v2</span></code> and <code class="docutils literal notranslate"><span class="pre">v3</span></code> is equal to <code class="docutils literal notranslate"><span class="pre">v4</span></code>.
Nothing is printed otherwise.
<textarea data-component="activecode" id=ifstmt3 data-lang="python" 
     data-include="ifstmt3_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your if statement
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Use <code class="docutils literal notranslate"><span class="pre">and</span></code> to join the two boolean expressions.</p>
<p>You do not need an else part.</p>
<p>Don’t forget to use <code class="docutils literal notranslate"><span class="pre">==</span></code>, not <code class="docutils literal notranslate"><span class="pre">=</span></code> (which is assignment).</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt3_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt3_a data-lang="python" 
     data-include="ifstmt3_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if v1 < v2 and v3 == v4:
    print('yes')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt4_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt4_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
v1 = 1
v2 = 2
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt4_q" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    <div id='ifstmt4_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt4" class="runestone explainer ac_section alert alert-warning">
If <code class="docutils literal notranslate"><span class="pre">v1</span></code> = 1, and <code class="docutils literal notranslate"><span class="pre">v2</span></code> = 2, create string variable <code class="docutils literal notranslate"><span class="pre">resStr</span></code> which refers to
<code class="docutils literal notranslate"><span class="pre">variable</span> <span class="pre">1</span> <span class="pre">is</span> <span class="pre">1</span> <span class="pre">(1.0)</span> <span class="pre">and</span> <span class="pre">variable</span> <span class="pre">2</span> <span class="pre">is</span> <span class="pre">2</span> <span class="pre">(2.0)</span></code>.
<textarea data-component="activecode" id=ifstmt4 data-lang="python" 
     data-include="ifstmt4_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(resStr, 'variable 1 is 1 (1.0) and variable 2 is 2 (2.0)')
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>To make sure you have only 1 digit after the decimal, use this format specifier: <code class="docutils literal notranslate"><span class="pre">%.1f</span></code></p>
<p>You can use a variable multiple times in the same tuple.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt4_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt4_a data-lang="python" 
     data-include="ifstmt4_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
resStr = 'variable 1 is %d (%.1f) and variable 2 is %d (%.1f)' % (v1, v1, v2, v2)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt5_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt5_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
v1 = 17
v4 = 4
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt5_q" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='ifstmt5_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt5" class="runestone explainer ac_section alert alert-warning">
Write an if statement that prints <code class="docutils literal notranslate"><span class="pre">is</span> <span class="pre">not</span> <span class="pre">3</span> <span class="pre">times</span></code> if <code class="docutils literal notranslate"><span class="pre">v1</span></code> is not equal to 3 times the value
of <code class="docutils literal notranslate"><span class="pre">v4</span></code>.
<textarea data-component="activecode" id=ifstmt5 data-lang="python" 
     data-include="ifstmt5_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>To make sure you have only 1 digit after the decimal, use this format specifier: <code class="docutils literal notranslate"><span class="pre">%.1f</span></code></p>
<p>Not equal to is <code class="docutils literal notranslate"><span class="pre">!=</span></code>.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt5_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt5_a data-lang="python" 
     data-include="ifstmt5_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if v1 != 3 * v4:
    print('is not 3 times')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt6_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt6_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
idx = 17
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt6_q" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    <div id='ifstmt6_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt6" class="runestone explainer ac_section alert alert-warning">
Write an if statement that assigns <code class="docutils literal notranslate"><span class="pre">res</span></code> to ‘even number’ if variable <code class="docutils literal notranslate"><span class="pre">idx</span></code>
is an even number and <code class="docutils literal notranslate"><span class="pre">res</span></code> to ‘odd number’ otherwise.
<textarea data-component="activecode" id=ifstmt6 data-lang="python" 
     data-include="ifstmt6_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(res, 'odd number')
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
You can tell if a value is even if when you divide it by 2, the remainder is 0.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt6_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt6_a data-lang="python" 
     data-include="ifstmt6_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if idx % 2 == 0:
    res = 'even number'
else:
    res = 'odd number'
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt7_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt7_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
v1 = 199.9999
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt7_q" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='ifstmt7_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt7" class="runestone explainer ac_section alert alert-warning">
Write an if statement that assigns <code class="docutils literal notranslate"><span class="pre">res</span></code> to 1 if variable <code class="docutils literal notranslate"><span class="pre">v1</span></code>
is strictly between 100 and 200.
<textarea data-component="activecode" id=ifstmt7 data-lang="python" 
     data-include="ifstmt7_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(res, 1)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Use <code class="docutils literal notranslate"><span class="pre">&lt;</span></code> and <code class="docutils literal notranslate"><span class="pre">&gt;</span></code> and <code class="docutils literal notranslate"><span class="pre">and</span></code>.</p>
<p>Remember that you have to have a value, variable, or expression on each side of
a <code class="docutils literal notranslate"><span class="pre">&lt;</span></code> or <code class="docutils literal notranslate"><span class="pre">&gt;</span></code>.  You cannot write <code class="docutils literal notranslate"><span class="pre">v1</span> <span class="pre">&gt;</span> <span class="pre">100</span> <span class="pre">and</span> <span class="pre">&lt;</span> <span class="pre">200</span></code> because there is no expression
on the left side of the <code class="docutils literal notranslate"><span class="pre">&lt;</span></code>.</p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt7_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt7_a data-lang="python" 
     data-include="ifstmt7_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if v1 > 100 and v1 < 200:
    res = 1
# could also write: if 100 < v1 < 200:
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt8_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt8_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
year = 2000
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt8_q" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    <div id='ifstmt8_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt8" class="runestone explainer ac_section alert alert-warning">
Write an if statement that assigns <code class="docutils literal notranslate"><span class="pre">leap</span></code> to <code class="docutils literal notranslate"><span class="pre">True</span></code> if variable <code class="docutils literal notranslate"><span class="pre">year</span></code> is a multiple of 4
and is not a multiple of 100, and to <code class="docutils literal notranslate"><span class="pre">False</span></code> otherwise.
<textarea data-component="activecode" id=ifstmt8 data-lang="python" 
     data-include="ifstmt8_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(leap, False)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
A number <code class="docutils literal notranslate"><span class="pre">n</span></code> is a multiple of another number <code class="docutils literal notranslate"><span class="pre">d</span></code>, if when you divide <code class="docutils literal notranslate"><span class="pre">n</span></code> by <code class="docutils literal notranslate"><span class="pre">d</span></code>,
the remainder is 0.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt8_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt8_a data-lang="python" 
     data-include="ifstmt8_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if year % 4 == 0 and year % 100 != 0:
    leap = True
else:
    leap = False
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt9_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt9_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
result = 44
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt9_q" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='ifstmt9_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt9" class="runestone explainer ac_section alert alert-warning">
Write an if statement that assigns 1 to <code class="docutils literal notranslate"><span class="pre">answer</span></code> if variable <code class="docutils literal notranslate"><span class="pre">result</span></code> has the value <code class="docutils literal notranslate"><span class="pre">100</span></code>,
and sets <code class="docutils literal notranslate"><span class="pre">answer</span></code> to 2 otherwise.
<textarea data-component="activecode" id=ifstmt9 data-lang="python" 
     data-include="ifstmt9_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(answer, 2)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use an if-else statement.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt9_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt9_a data-lang="python" 
     data-include="ifstmt9_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if result == 100:
    answer = 1
else:
    answer = 2
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="ifstmt10_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt10_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
score = 72
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="ifstmt10_q" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    <div id='ifstmt10_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt10" class="runestone explainer ac_section alert alert-warning">
Write an ifelse statement that sets the value of a variable <code class="docutils literal notranslate"><span class="pre">gradepoint</span></code> to 4
if a variable <code class="docutils literal notranslate"><span class="pre">score</span></code> is greater than 90, 3 if <code class="docutils literal notranslate"><span class="pre">score</span></code> is between 80 and 89,
2 if <code class="docutils literal notranslate"><span class="pre">score</span></code> is between 70 and 79, 1 if <code class="docutils literal notranslate"><span class="pre">score</span></code> is between 60 and 69, and 0 otherwise.
<textarea data-component="activecode" id=ifstmt10 data-lang="python" 
     data-include="ifstmt10_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
   def testOne(self):
       self.assertEqual(gradepoint, 2)
myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use an if-elif-else statement.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt10_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt10_a data-lang="python" 
     data-include="ifstmt10_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
if score > 90:
    gradepoint = 4
elif score > 80:
    gradepoint = 3
elif score > 70:
    gradepoint = 2
elif score > 60:
    gradepoint = 1
else:
    gradepoint = 0
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="ifstmt11_q" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    <div id='ifstmt11_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt11" class="runestone explainer ac_section alert alert-warning">
Write a short program that first asks the user to input 3 numbers.
(Do this with 3 individual <code class="docutils literal notranslate"><span class="pre">input()</span></code> calls, storing the results in 3 variables.)
Then, the program prints out the largest of the 3 numbers.
You may assume the user enters three distinct values.
<textarea data-component="activecode" id=ifstmt11 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use multiple if/elif statements, each with multiple boolean operators connected together with <code class="docutils literal notranslate"><span class="pre">and</span></code> s.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt11_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt11_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
v1 = float(input("Enter first value:"))
v2 = float(input("Enter second value:"))
v3 = float(input("Enter third value:"))
if v1 > v2 and v1 > v3:
    print(v1)
elif v2 > v1 and v2 > v3:
    print(v2)
else:
    print(v3)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="ifstmt12_q" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    <div id='ifstmt12_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt12" class="runestone explainer ac_section alert alert-warning">
Write a short program that asks the user for a payrate and a number of hours, and then
prints out the total pay. Any hours over 40 are paid with time and a half.
<textarea data-component="activecode" id=ifstmt12 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use an if-else statement.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt12_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt12_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
payrate = float(input("Enter hourly payrate: "))
hours = float(input("Enter total hours worked: "))
if hours > 40:
    pay = payrate * 1.5 * (hours - 40) + payrate * 40
else:
    pay = payrate * hours
print('Your pay will is', pay)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="ifstmt13_q" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    <div id='ifstmt13_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt13" class="runestone explainer ac_section alert alert-warning">
<p>On a roulette wheel, the pockets are numbered from 0 to 36.  The colors of the pockets
are as follows:</p>
<ul class="simple">
<li>Pocket 0 is green.</li>
<li>For pockets 1 through 10, the odd-numbered pockets are red and the even-numbered pockets are black.</li>
<li>For pockets 11 through 18, the odd-numbered pockets are black and the even-numbered pockets are red.</li>
<li>For pockets 19 through 28, the odd-numbered pockets are red and the even-numbered pockets are black.</li>
<li>For pockets 29 through 36, the odd-numbered pockets are black and the even-numbered pockets are red.</li>
</ul>
<p>Write a program that asks the user to enter a pocket number
and displays whether the pocket is green, red, or black.
The program should display an error message if the user
enters a number that is outside the range of 0 through 36.</p>

<textarea data-component="activecode" id=ifstmt13 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use an if-elif-else statement with nested if statements in each
part.  An odd number is one where when you divide by 2, the
remainder is 1.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt13_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt13_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
pocket = int(input("Enter a pocket number from 0 to 36, inclusive: "))
if pocket < 0 or pocket > 36:
    print("That is not a legal pocket number.")
elif pocket == 0:
    print("Pocket is green")
elif pocket >= 1 and pocket <= 10:
    if pocket % 2 == 1:
        print("Pocket is red")
    else:
        print("Pocket is black")
elif pocket >= 11 and pocket <= 18:
    if pocket % 2 == 0:
        print("Pocket is red")
    else:
        print("Pocket is black")
elif pocket >= 19 and pocket <= 28:
    if pocket % 2 == 1:
        print("Pocket is red")
    else:
        print("Pocket is black")
else:
    if pocket % 2 == 0:
        print("Pocket is red")
    else:
        print("Pocket is black")
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="ifstmt14_q" >
    <ol start=14 class=arabic><li class="alert alert-warning">

    <div id='ifstmt14_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="ifstmt14" class="runestone explainer ac_section alert alert-warning">
<p>Write a program that asks the user to enter two numbers and
an operator to apply to them.  The operator must be one of
<code class="docutils literal notranslate"><span class="pre">+</span></code>, <code class="docutils literal notranslate"><span class="pre">-</span></code>, <code class="docutils literal notranslate"><span class="pre">*</span></code>, <code class="docutils literal notranslate"><span class="pre">/</span></code>, <code class="docutils literal notranslate"><span class="pre">//</span></code>, or <code class="docutils literal notranslate"><span class="pre">%</span></code>. The program
applies the operator to the numbers and prints the result.</p>
<p>Here is an example run:</p>
<p>Enter first number: <code class="docutils literal notranslate"><span class="pre">46</span></code></p>
<p>Enter second number: <code class="docutils literal notranslate"><span class="pre">13.5</span></code></p>
<p>Enter an operator (one of +, -, *, /, //, %): <code class="docutils literal notranslate"><span class="pre">-</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">46.0000000</span> <span class="pre">-</span> <span class="pre">13.5000000</span> <span class="pre">=</span> <span class="pre">32.500000</span></code></p>
<p>If the user enters an illegal operator, the program prints
an error message.</p>

<textarea data-component="activecode" id=ifstmt14 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<p>Use an if-elif-else statement.
The print statement for the + operator looks like this:</p>
<p><code class="docutils literal notranslate"><span class="pre">print('%f</span> <span class="pre">+</span> <span class="pre">%f</span> <span class="pre">=</span> <span class="pre">%f'</span> <span class="pre">%</span> <span class="pre">(num1,</span> <span class="pre">num2,</span> <span class="pre">num1</span> <span class="pre">+</span> <span class="pre">num2))</span></code></p>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="ifstmt14_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ifstmt14_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
num1 = float(input("Enter first number: "))
num2 = float(input("Enter second number: "))
op = input("Enter an operator (one of +, -, *, /, //, %): ")
if op == '+':
    print('%f + %f = %f' % (num1, num2, num1 + num2))
elif op == '-':
    print('%f - %f = %f' % (num1, num2, num1 - num2))
elif op == '*':
    print('%f * %f = %f' % (num1, num2, num1 * num2))
elif op == '/':
    print('%f / %f = %f' % (num1, num2, num1 / num2))
elif op == '//':
    print('%f // %f = %f' % (num1, num2, num1 // num2))
elif op == '%':
    print('%f % %f = %f' % (num1, num2, num1 % num2))
else:
    print('%s is an unknown operator' % op)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    </div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - String Formating' data-toggle="tooltip">
    <a href="string_formatting.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="string_formatting.html" ></a>
  

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2019 vtn2.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 3.2.15.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>