<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Python Problem Sets &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="_static/video.css" />
    <link rel="stylesheet" type="text/css" href="_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="_static/reveal.css" />
    <link rel="stylesheet" type="text/css" href="_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="_static/js_lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="_static/accessibility.css" />
    <link rel="stylesheet" type="text/css" href="_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="_static/poll.css" />
    <link rel="stylesheet" href="_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="./" src="_static/documentation_options.js"></script>
    <script type="text/javascript" src="_static/runestonebase.js"></script>
    <script type="text/javascript" src="_static/jquery.js"></script>
    <script type="text/javascript" src="_static/underscore.js"></script>
    <script type="text/javascript" src="_static/doctools.js"></script>
    <script type="text/javascript" src="_static/matrixeq.js"></script>
    <script type="text/javascript" src="_static/clickable.js"></script>
    <script type="text/javascript" src="_static/timedclickable.js"></script>
    <script type="text/javascript" src="_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="_static/fitb.js"></script>
    <script type="text/javascript" src="_static/timedfitb.js"></script>
    <script type="text/javascript" src="_static/runestonevideo.js"></script>
    <script type="text/javascript" src="_static/codemirror.js"></script>
    <script type="text/javascript" src="_static/xml.js"></script>
    <script type="text/javascript" src="_static/css.js"></script>
    <script type="text/javascript" src="_static/htmlmixed.js"></script>
    <script type="text/javascript" src="_static/javascript.js"></script>
    <script type="text/javascript" src="_static/webglinteractive.js"></script>
    <script type="text/javascript" src="_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="_static/Blob.js"></script>
    <script type="text/javascript" src="_static/reveal.js"></script>
    <script type="text/javascript" src="_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="_static/parsons.js"></script>
    <script type="text/javascript" src="_static/timedparsons.js"></script>
    <script type="text/javascript" src="_static/animationbase.js"></script>
    <script type="text/javascript" src="_static/mchoice.js"></script>
    <script type="text/javascript" src="_static/timedmc.js"></script>
    <script type="text/javascript" src="_static/timed.js"></script>
    <script type="text/javascript" src="_static/shortanswer.js"></script>
    <script type="text/javascript" src="_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="_static/pytutor.js"></script>
    <script type="text/javascript" src="_static/codelens.js"></script>
    <script type="text/javascript" src="_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="_static/bookfuncs.js"></script>
    <script type="text/javascript" src="_static/python.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="_static/skulpt.min.js"></script>
    <script type="text/javascript" src="_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="_static/activecode.js"></script>
    <script type="text/javascript" src="_static/clike.js"></script>
    <script type="text/javascript" src="_static/timed_activecode.js"></script>
    <script type="text/javascript" src="_static/datafile.js"></script>
    <script type="text/javascript" src="_static/showEval.js"></script>
    <script type="text/javascript" src="_static/dragndrop.js"></script>
    <script type="text/javascript" src="_static/timeddnd.js"></script>
    <script type="text/javascript" src="_static/poll.js"></script>
    <script type="text/javascript" src="_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="_static/jquery-fix.js"></script>
    <script type="text/javascript" src="_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="_static/waypoints.min.js"></script>
    <script type="text/javascript" src="_static/user-highlights.js"></script>
    <script type="text/javascript" src="_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="genindex.html" />
    <link rel="search" title="Search" href="search.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'python_problems';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = false;
  eBookConfig.basecourse = 'python_problems';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '3.2.15';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="#">Problems in Python</a>
      </div>
    </div>

     

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='#'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/python_problems/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/python_problems/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=python_problems&page=index'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Python Problem Sets</a><ul>
<li><a class="reference internal" href="#section-1-simple-assignments">SECTION 1: Simple Assignments</a></li>
<li><a class="reference internal" href="#section-2-input-and-output">SECTION 2: Input and Output</a></li>
<li><a class="reference internal" href="#section-1-introduction">SECTION 1: Introduction</a></li>
<li><a class="reference internal" href="#section-2-links">Section 2: Links</a></li>
<li><a class="reference internal" href="#section-3-sample-directives">SECTION 3: Sample Directives</a><ul>
<li><a class="reference internal" href="#activecode">ActiveCode</a></li>
<li><a class="reference internal" href="#multiple-choice">Multiple Choice</a></li>
</ul>
</li>
</ul>
</li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>
      
    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="python-problem-sets">
<h1>Python Problem Sets<a class="headerlink" href="#python-problem-sets" title="Permalink to this headline">¶</a></h1>
<div class="section" id="section-1-simple-assignments">
<h2>SECTION 1: Simple Assignments<a class="headerlink" href="#section-1-simple-assignments" title="Permalink to this headline">¶</a></h2>

    <div data-component="question" class="full-width container question" id="sass1_q" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <div id='sass1_tab' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass1" class="runestone explainer ac_section alert alert-warning">
Write code to create a variable <code class="docutils literal notranslate"><span class="pre">score</span></code> and make it refer to
the integer 1.
<textarea data-component="activecode" id=sass1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(score, 1)
      self.assertIs(type(score), int, "variable must refer to an int")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Remember that an assignment is a <code class="docutils literal notranslate"><span class="pre">=</span></code> with a variable on the left-hand
side and a value or expression on the right-hand side.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass1_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass1_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
score = 1
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="sass2_q" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    <div id='sass2_tab' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass2" class="runestone explainer ac_section alert alert-warning">
Write code to create a variable <code class="docutils literal notranslate"><span class="pre">temperature</span></code> and make it refer to a floating point number.
<textarea data-component="activecode" id=sass2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertIs(type(temperature), float, "variable must refer to a float")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
A floating point number is a real number. I.e., it has a <code class="docutils literal notranslate"><span class="pre">.</span></code> in it.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass2_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass2_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
temperature = 57.3
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="sass3_q" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='sass3_tab' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass3" class="runestone explainer ac_section alert alert-warning">
Write code to create a variable <code class="docutils literal notranslate"><span class="pre">name</span></code> and make it refer to a string.
<textarea data-component="activecode" id=sass3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertIs(type(name), str, "variable must refer to a string")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
A string is enclosed by single or double quotes.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass3_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass3_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
name = "Desmond Tutu"
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="sass4_q" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    <div id='sass4_tab' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass4" class="runestone explainer ac_section alert alert-warning">
Write code to create a variable <code class="docutils literal notranslate"><span class="pre">matches</span></code> and make it refer to a boolean value.
<textarea data-component="activecode" id=sass4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertIs(type(matches), bool, "variable must refer to a boolean")

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
The two booleans values are <code class="docutils literal notranslate"><span class="pre">True</span></code> and <code class="docutils literal notranslate"><span class="pre">False</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass4_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass4_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
matches = False
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="sass5_q" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='sass5_tab' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass5" class="runestone explainer ac_section alert alert-warning">
Assume the variable <code class="docutils literal notranslate"><span class="pre">temperature</span></code> has been initialized to some value.
Write code to add 1 to the variable <code class="docutils literal notranslate"><span class="pre">temperature</span></code>.
<textarea data-component="activecode" id=sass5 data-lang="python" 
     data-include="sass2_a" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(temperature, 58.3)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
The right-hand side of an equation is evaluated first, so you can take
<code class="docutils literal notranslate"><span class="pre">temperature</span></code> and add 1 to it on the right-hand side, and then put
<code class="docutils literal notranslate"><span class="pre">temperature</span></code> on the left-hand side of the <code class="docutils literal notranslate"><span class="pre">=</span></code> sign.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass5_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass5_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
temperature = temperature + 1
# Also, could be written temperature += 1
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="sass6_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass6_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
hours = 30.0
hourly_rate = 11.50
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="sass6_q" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    <div id='sass6_tab' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass6" class="runestone explainer ac_section alert alert-warning">
Assume you have two variables <code class="docutils literal notranslate"><span class="pre">hours</span></code> and <code class="docutils literal notranslate"><span class="pre">hourly_rate</span></code>.  Write code to compute the
<code class="docutils literal notranslate"><span class="pre">total_pay</span></code>, by multiplying <code class="docutils literal notranslate"><span class="pre">hours</span></code> and <code class="docutils literal notranslate"><span class="pre">hourly_rate</span></code> and storing in <code class="docutils literal notranslate"><span class="pre">total_pay</span></code>.
<textarea data-component="activecode" id=sass6 data-lang="python" 
     data-include="sass6_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# Assume hours and hourly_rate have been defined and given values.

# replace this comment with your code
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(total_pay, hours * hourly_rate)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Multiply <code class="docutils literal notranslate"><span class="pre">hours</span></code> and <code class="docutils literal notranslate"><span class="pre">hourly_rate</span></code> on the right-hand side of the assignment statement.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass6_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass6_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
total_pay = hours * hourly_rate
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="sass7_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass7_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
p1_points = 17
p2_points = 44
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="sass7_q" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='sass7_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass7" class="runestone explainer ac_section alert alert-warning">
Assume you have two variables <code class="docutils literal notranslate"><span class="pre">p1_points</span></code> and <code class="docutils literal notranslate"><span class="pre">p2_points</span></code>.  Write a line of code
to indicate that player p1 added all of player p2’s points to her own.
<textarea data-component="activecode" id=sass7 data-lang="python" 
     data-include="sass7_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(p1_points, 61)
        self.assertAlmostEqual(p2_points, 44)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Your code needs to add <code class="docutils literal notranslate"><span class="pre">p2_points</span></code> to <code class="docutils literal notranslate"><span class="pre">p1_points</span></code> and make <code class="docutils literal notranslate"><span class="pre">p1_points</span></code> refer to the
result.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass7_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass7_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
p1_points = p1_points + p2_points
# or  p1_points = p2_points + p1_points
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="sass8_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass8_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
ball1_direction = 17
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="sass8_q" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    <div id='sass8_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="sass8" class="runestone explainer ac_section alert alert-warning">
Assume you have a variable <code class="docutils literal notranslate"><span class="pre">ball1_direction</span></code>.  Write a line of code
that set ball2’s direction to be the same as ball1’s direction.
<textarea data-component="activecode" id=sass8 data-lang="python" 
     data-include="sass8_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# replace this comment with your code
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(ball2_direction, ball1_direction)
        self.assertAlmostEqual(ball1_direction, 17)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Your code creates variable <code class="docutils literal notranslate"><span class="pre">ball2_direction</span></code> on the left-hand side of the <code class="docutils literal notranslate"><span class="pre">=</span></code>
sign and sets it value to <code class="docutils literal notranslate"><span class="pre">ball1_direction</span></code>.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass8_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass8_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
ball2_direction = ball1_direction
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="sass9_q" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='sass9_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Replace these statements with a single statement so that you don’t use the variable <code class="docutils literal notranslate"><span class="pre">y</span></code> –
just the variable <code class="docutils literal notranslate"><span class="pre">x</span></code> being set to a value.</p>

<div data-childcomponent="sass9" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass9 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
x = 7
y = x + 1
x = y
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(x, 8)

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="sass9_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass9_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
     x = 8
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="sass10_q" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    <div id='sass10_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >
<p>Before running the following code, predict what value will be printed.</p>

<div data-childcomponent="sass10" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=sass10 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
x = 7
y = 9
y = x
x = 4
print(y)
</textarea>
</div>
<p>If you don’t understand the answer, using <strong>CodeLens</strong> might help.</p>
</div>
    </div>


    </li></ol>
    </div>
    </div>
<div class="section" id="section-2-input-and-output">
<h2>SECTION 2: Input and Output<a class="headerlink" href="#section-2-input-and-output" title="Permalink to this headline">¶</a></h2>

    <div data-component="question" class="full-width container question" id="io1_q" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <div id='io1_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io1" class="runestone explainer ac_section alert alert-warning">
Write code to print <code class="docutils literal notranslate"><span class="pre">Hello,</span> <span class="pre">World</span></code>
<textarea data-component="activecode" id=io1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use the print() function, and remember that strings are surrounded by double-quotes or single-quotes.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io1_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io1_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print("Hello, World")
# or print('Hello, World')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io2_q" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    <div id='io2_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io2" class="runestone explainer ac_section alert alert-warning">
<p>Write code to print</p>
<blockquote>
<div><p><code class="docutils literal notranslate"><span class="pre">Hello</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">World</span></code></p>
</div></blockquote>
<p>using 2 print statements.</p>

<textarea data-component="activecode" id=io2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use the print() function, and remember that strings are surrounded by double-quotes or single-quotes.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io2_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io2_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print("Hello")
print("World")
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <p>Question 3:</p>
<p>Write code to print <strong>Hello</strong> leaving the cursor on the same line.</p>
<p>Question 4:</p>
<p>Assume you have a variable <code class="docutils literal notranslate"><span class="pre">ranking</span></code> set to some integer value. Write a line of code to print the value of <code class="docutils literal notranslate"><span class="pre">ranking</span></code>.</p>
<p>Question 5:</p>
<p>Assume you have two variables <code class="docutils literal notranslate"><span class="pre">ranking</span></code> and <code class="docutils literal notranslate"><span class="pre">average</span></code> set to come values. Write a line of
code to print the values with a single space between.</p>
<p>Question 6:</p>
<p>Assume you have two variables <code class="docutils literal notranslate"><span class="pre">start</span></code> and <code class="docutils literal notranslate"><span class="pre">interval</span></code>. Write a line of code to print the sum of
of the two values.</p>
<p>Question 7:</p>
<p>Assume you have a variable <code class="docutils literal notranslate"><span class="pre">ranking</span></code> set to some integer value.  Write a line of code to print <strong>Ranking :</strong> followed
by the value <code class="docutils literal notranslate"><span class="pre">ranking</span></code> refers to.  E.g., if <code class="docutils literal notranslate"><span class="pre">ranking</span></code> was 7, the output would be <strong>Ranking: 7</strong>.</p>
<p>Question 8:</p>
<p>Assume you have</p>
</div>
<div class="section" id="section-1-introduction">
<h2>SECTION 1: Introduction<a class="headerlink" href="#section-1-introduction" title="Permalink to this headline">¶</a></h2>
<p>Congratulations!   If you can see this file you have probably successfully run the <code class="docutils literal notranslate"><span class="pre">runestone</span> <span class="pre">init</span></code> command.  If you are looking at this as a source file you should now run <code class="docutils literal notranslate"><span class="pre">runestone</span> <span class="pre">build</span></code>  to generate html files.   Once you have run the build command you can run <code class="docutils literal notranslate"><span class="pre">runestone</span> <span class="pre">serve</span></code> and then view this in your browser at <code class="docutils literal notranslate"><span class="pre">http://localhost:8000</span></code></p>
<p>This is just a sample of what you can do.  The index.rst file is the table of contents for your entire project.  You can put all of your writing in the index, or  you can include additional rst files.  Those files may even be in subdirectories that you can reference using a relative path.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="o">..</span> <span class="n">toctree</span><span class="p">::</span>
   <span class="p">:</span><span class="n">maxdepth</span><span class="p">:</span> <span class="mi">2</span>

   <span class="n">some</span><span class="o">/</span><span class="n">path</span><span class="o">/</span><span class="n">myfile</span><span class="o">.</span><span class="n">rst</span>
</pre></div>
</div>
</div>
<div class="section" id="section-2-links">
<h2>Section 2: Links<a class="headerlink" href="#section-2-links" title="Permalink to this headline">¶</a></h2>
<p>Runestone uses the <code class="docutils literal notranslate"><span class="pre">restructuredText</span></code> (rst) markup language.  We chose this over markdown largely because rst is extensible.  Nearly all of the basic markup tasks are already handled by restructuredText.  You should check out the docs for the basics of restructuredText (link below). Our extensions are all for the interactive elements.  One key hint about restructuredText:  Its like <strong>Python</strong> – <em>indentation matters!</em></p>
<ul class="simple">
<li><a class="reference external" href="http://docutils.sourceforge.net/rst.html">restructuredText Docs</a></li>
<li><a class="reference external" href="http://runestoneinteractive.org/build/html/index.html">Runestone Docs</a></li>
<li>Join the discussion on our <a class="reference external" href="https://groups.google.com/forum/#!forum/runestone_instructors">Google Group</a></li>
<li>Tell us about problems on <a class="reference external" href="https://github.com/RunestoneInteractive/RunestoneComponents">Github</a></li>
</ul>
</div>
<div class="section" id="section-3-sample-directives">
<h2>SECTION 3: Sample Directives<a class="headerlink" href="#section-3-sample-directives" title="Permalink to this headline">¶</a></h2>
<div class="section" id="activecode">
<h3>ActiveCode<a class="headerlink" href="#activecode" title="Permalink to this headline">¶</a></h3>

<div data-childcomponent="codeexample1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=codeexample1 data-lang="python" 
      data-timelimit=25000 data-coach="true" data-codelens="true"  
    data-audio=''    
         data-caption='This is a caption' >
print("My first program adds a list of numbers")
myList = [2, 4, 6, 8, 10]
total = 0
for num in myList:
    total = total + num
print(total)
</textarea>
</div>
</div>
<div class="section" id="multiple-choice">
<h3>Multiple Choice<a class="headerlink" href="#multiple-choice" title="Permalink to this headline">¶</a></h3>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-multipleanswers="true"  id="question1_2">
            Q-13: Which colors might be found in a rainbow? (choose all that are correct)
            <li data-component="answer" data-correct id="question1_2_opt_a">red</li><li data-component="feedback" id="question1_2_opt_a">Red is a definitely on of the colors.</li>
            
            <li data-component="answer" data-correct id="question1_2_opt_b">yellow</li><li data-component="feedback" id="question1_2_opt_b">Yes, yellow is correct.</li>
            
            <li data-component="answer"  id="question1_2_opt_c">black</li><li data-component="feedback" id="question1_2_opt_c">Remember the acronym...ROY G BIV.  B stands for blue.</li>
            
            <li data-component="answer" data-correct id="question1_2_opt_d">green</li><li data-component="feedback" id="question1_2_opt_d">Yes, green is one of the colors.</li>
            

            </ul>
            </div>
            <p>These are just two of the many interactive components for writing online course materials.  You can see examples of all of them <a class="reference external" href="http://interactivepython.org/runestone/static/overview/overview.html">On our Example Page</a></p>
<p>Now feel free to modify this file to start creating your own interactive page.</p>
</div>
</div>
</div>


  
      

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2017 vtn2.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 3.2.15.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>