<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>Input and Output &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/presenter_mode.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="_static/video.css" />
    <link rel="stylesheet" type="text/css" href="_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="_static/reveal.css" />
    <link rel="stylesheet" type="text/css" href="_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="_static/js_lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="_static/accessibility.css" />
    <link rel="stylesheet" type="text/css" href="_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="_static/poll.css" />
    <link rel="stylesheet" href="_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="./" src="_static/documentation_options.js"></script>
    <script type="text/javascript" src="_static/runestonebase.js"></script>
    <script type="text/javascript" src="_static/jquery.js"></script>
    <script type="text/javascript" src="_static/underscore.js"></script>
    <script type="text/javascript" src="_static/doctools.js"></script>
    <script type="text/javascript" src="_static/matrixeq.js"></script>
    <script type="text/javascript" src="_static/clickable.js"></script>
    <script type="text/javascript" src="_static/timedclickable.js"></script>
    <script type="text/javascript" src="_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="_static/fitb.js"></script>
    <script type="text/javascript" src="_static/timedfitb.js"></script>
    <script type="text/javascript" src="_static/runestonevideo.js"></script>
    <script type="text/javascript" src="_static/codemirror.js"></script>
    <script type="text/javascript" src="_static/xml.js"></script>
    <script type="text/javascript" src="_static/css.js"></script>
    <script type="text/javascript" src="_static/htmlmixed.js"></script>
    <script type="text/javascript" src="_static/javascript.js"></script>
    <script type="text/javascript" src="_static/webglinteractive.js"></script>
    <script type="text/javascript" src="_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="_static/Blob.js"></script>
    <script type="text/javascript" src="_static/reveal.js"></script>
    <script type="text/javascript" src="_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="_static/parsons.js"></script>
    <script type="text/javascript" src="_static/timedparsons.js"></script>
    <script type="text/javascript" src="_static/animationbase.js"></script>
    <script type="text/javascript" src="_static/mchoice.js"></script>
    <script type="text/javascript" src="_static/timedmc.js"></script>
    <script type="text/javascript" src="_static/timed.js"></script>
    <script type="text/javascript" src="_static/shortanswer.js"></script>
    <script type="text/javascript" src="_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="_static/pytutor.js"></script>
    <script type="text/javascript" src="_static/codelens.js"></script>
    <script type="text/javascript" src="_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="_static/bookfuncs.js"></script>
    <script type="text/javascript" src="_static/python.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="_static/skulpt.min.js"></script>
    <script type="text/javascript" src="_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="_static/activecode.js"></script>
    <script type="text/javascript" src="_static/clike.js"></script>
    <script type="text/javascript" src="_static/timed_activecode.js"></script>
    <script type="text/javascript" src="_static/datafile.js"></script>
    <script type="text/javascript" src="_static/showEval.js"></script>
    <script type="text/javascript" src="_static/dragndrop.js"></script>
    <script type="text/javascript" src="_static/timeddnd.js"></script>
    <script type="text/javascript" src="_static/poll.js"></script>
    <script type="text/javascript" src="_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="_static/jquery-fix.js"></script>
    <script type="text/javascript" src="_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="_static/waypoints.min.js"></script>
    <script type="text/javascript" src="_static/user-highlights.js"></script>
    <script type="text/javascript" src="_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="genindex.html" />
    <link rel="search" title="Search" href="search.html" />
    <link rel="next" title="Expressions" href="expressions.html" />
    <link rel="prev" title="Simple Assignments" href="assignments.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'python_problems';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'python_problems';
  eBookConfig.acDefaultLanguage = 'python3' ? 'python3' : 'python'
  eBookConfig.runestone_version = '3.2.15';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="index.html">Problems in Python</a>
      </div>
    </div>

     

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/python_problems/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/python_problems/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=python_problems&page=io'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>
      </div>






      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Input and Output</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Simple Assignments" data-toggle="tooltip">
    <a href="assignments.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Expressions' data-toggle="tooltip" >              
    <a href="expressions.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>
      
    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="input-and-output">
<h1>Input and Output<a class="headerlink" href="#input-and-output" title="Permalink to this headline">¶</a></h1>

    <div data-component="question" class="full-width container question" id="io1_q" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    <div id='io1_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io1" class="runestone explainer ac_section alert alert-warning">
Write code to print <code class="docutils literal notranslate"><span class="pre">Hello,</span> <span class="pre">World</span></code>
<textarea data-component="activecode" id=io1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use the print() function, and remember that strings are surrounded by double-quotes or single-quotes.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io1_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io1_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print("Hello, World")
# or print('Hello, World')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io2_q" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    <div id='io2_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io2" class="runestone explainer ac_section alert alert-warning">
<p>Write code to print</p>
<blockquote>
<div><p><code class="docutils literal notranslate"><span class="pre">Hello</span></code></p>
<p><code class="docutils literal notranslate"><span class="pre">World</span></code></p>
</div></blockquote>
<p>using 2 print statements.</p>

<textarea data-component="activecode" id=io2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Call <code class="docutils literal notranslate"><span class="pre">print()</span></code> twice, on two consecutive lines.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io2_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io2_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print("Hello")
print("World")
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io3_q" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    <div id='io3_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io3" class="runestone explainer ac_section alert alert-warning">
<p>Write code to print</p>
<blockquote>
<div><code class="docutils literal notranslate"><span class="pre">Hello</span></code></div></blockquote>
<p>leaving the cursor on the same line.</p>

<textarea data-component="activecode" id=io3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >

</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
You have to use the optional parameter, <code class="docutils literal notranslate"><span class="pre">end</span></code>, in the print statement.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io3_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io3_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print("Hello", end='')
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io4_q" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    <div id='io4_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io4" class="runestone explainer ac_section alert alert-warning">
Assume you have a variable <code class="docutils literal notranslate"><span class="pre">ranking</span></code> set to some integer value.
Write a line of code to print the value of <code class="docutils literal notranslate"><span class="pre">ranking</span></code>.
<textarea data-component="activecode" id=io4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
ranking = 99
# Replace this comment with your code
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
<code class="docutils literal notranslate"><span class="pre">print()</span></code> evaluates each variable or expression before printing it.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io4_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io4_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print(ranking)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="io5_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io5_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
ranking = 32
average = 34
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="io5_q" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    <div id='io5_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io5" class="runestone explainer ac_section alert alert-warning">
Assume you have two variables <code class="docutils literal notranslate"><span class="pre">ranking</span></code> and <code class="docutils literal notranslate"><span class="pre">average</span></code> set to some values.
Write a line of code to print the values with a single space between.
<textarea data-component="activecode" id=io5 data-lang="python" 
     data-include="io5_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# Replace this comment with your code.
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
The comma in <code class="docutils literal notranslate"><span class="pre">print('x',</span> <span class="pre">'y')</span></code> will automatically insert a space between the two values</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io5_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io5_a data-lang="python" 
     data-include="io5_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print(ranking, average)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io6_q" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    <div id='io6_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io6" class="runestone explainer ac_section alert alert-warning">
Assume you have two variables <code class="docutils literal notranslate"><span class="pre">start</span></code> and <code class="docutils literal notranslate"><span class="pre">interval</span></code>. Write a line of
code to print the sum of the two values.
<textarea data-component="activecode" id=io6 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
start = 103233.1
interval = 201787.33
# Replace this comment with your code.
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
You can put expressions, like <code class="docutils literal notranslate"><span class="pre">x</span> <span class="pre">+</span> <span class="pre">y</span></code> into a <code class="docutils literal notranslate"><span class="pre">print</span></code> statement, too.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io6_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io6_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print(start + interval)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io7_q" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    <div id='io7_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io7" class="runestone explainer ac_section alert alert-warning">
Assume you have a variable <code class="docutils literal notranslate"><span class="pre">ranking</span></code> set to some integer value.
Write a line of code to print <code class="docutils literal notranslate"><span class="pre">Ranking:</span></code> followed by
the value that <code class="docutils literal notranslate"><span class="pre">ranking</span></code> refers to.  Note that there should be
one space between the <code class="docutils literal notranslate"><span class="pre">:</span></code> and the value.
E.g., if <code class="docutils literal notranslate"><span class="pre">ranking</span></code> was 7, the output would be <code class="docutils literal notranslate"><span class="pre">Ranking:</span> <span class="pre">7</span></code>.
<textarea data-component="activecode" id=io7 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
ranking = 7
# Replace this comment with your code.
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Remember, the comma in a <code class="docutils literal notranslate"><span class="pre">print</span></code> adds a space, so be careful!</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io7_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io7_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print('Ranking:', ranking)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    <div style='display:none;'>
<div data-childcomponent="io8_pre" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io8_pre data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
name = 'Joe McDonald'
score = 7
</textarea>
</div>
</div>
    <div data-component="question" class="full-width container question" id="io8_q" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    <div id='io8_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io8" class="runestone explainer ac_section alert alert-warning">
Assume you have a variable <code class="docutils literal notranslate"><span class="pre">name</span></code> set to a student’s name, and
a variable <code class="docutils literal notranslate"><span class="pre">score</span></code> set to the score the student got on the last quiz.
Write a line of code to print out the student’s name and the student’s score,
in the following format.
Name: Joe McDonald Score: 7
where the student’s name is <cite>Joe McDonald</cite> and the student’s score is <cite>7</cite>.
<textarea data-component="activecode" id=io8 data-lang="python" 
     data-include="io8_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# Replace this comment with your code.
====
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
You can pass as many parameters as you want to <code class="docutils literal notranslate"><span class="pre">print</span></code>, separated by commas.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io8_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io8_a data-lang="python" 
     data-include="io8_pre" data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
print('Name:', name, 'Score:', score)
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io9_q" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    <div id='io9_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io9" class="runestone explainer ac_section alert alert-warning">
Write code to read in a string value from the user, storing the result
in a variable <code class="docutils literal notranslate"><span class="pre">name</span></code>.
<textarea data-component="activecode" id=io9 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# Replace this comment with your code.
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
       self.assertIs(type(name), str, 'variable must be a str')

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use <code class="docutils literal notranslate"><span class="pre">input()</span></code>.  You do <em>NOT</em> have to convert the result to a string, because
<code class="docutils literal notranslate"><span class="pre">input()</span></code> returns a string.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io9_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io9_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
name = input()
# or, better: name = input("Enter a name: ")
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="io10_q" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    <div id='io10_tabs' data-component="tabbedStuff"  class='alert alert-warning'><div data-component="tab" data-tabname="Question" >

<div data-childcomponent="io10" class="runestone explainer ac_section alert alert-warning">
Write code to prompt the user for a score, reading in the score into a variable
<code class="docutils literal notranslate"><span class="pre">score</span></code>.  The type of score should be a <code class="docutils literal notranslate"><span class="pre">float</span></code>.
<textarea data-component="activecode" id=io10 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
# Replace this comment with your code.
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
       self.assertIs(type(score), float, 'variable must be a float')

myTests().main()
</textarea>
</div>
</div><div data-component="tab" data-tabname="Hint" >
Use <code class="docutils literal notranslate"><span class="pre">input()</span></code>.  You have to convert the result to a float, because
<code class="docutils literal notranslate"><span class="pre">input()</span></code> returns a string.</div><div data-component="tab" data-tabname="Answer" >

<div data-childcomponent="io10_a" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=io10_a data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''    
          >
score = float(input('Enter a score: '))
# or,
# print("Enter a score: ", end='')
# score = float(input())
</textarea>
</div>
</div>
    </div>


    </li></ol>
    </div>
    </div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Simple Assignments' data-toggle="tooltip">
    <a href="assignments.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="assignments.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Expressions' data-toggle="tooltip" >
    <a href="expressions.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="expressions.html" >Next Section - Expressions</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2019 vtn2.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 3.2.15.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>